# Erste Schritte mit Colab und Python


Tradition gebietet es, mit *Hello world* anzufangen. Einzige Besonderheit hier: Strings in python können mit doppelten oder einfachen Anführungszeichen definiert werden -- macht keinen Unterschied. 

```python
print("Hello world!")
print('Hello world!')
```

Im Gegensatz zu *java* oder *C*, muss der Typ von Variablen nicht deklariert werden. Dh, wenn wir eine Variable verwenden wollen, so weisen wir ihr einfach einen Wert zu und können sie dann verwenden. 

```python
x=42
print(x)
```

Hier ist eine Variable, die einen String enthält:

```python
s='some string'
s
```

Den Typ einer Variable erhalten wir mit dem eingebauten Befehl <code>type</code>:

```python
type(x),type(s)
```

Weitere Typen von Variable sind <code>float</code> (Fließkommazahlen, also 1.2345), <code>bool</code> (*boolean*, also entweder <code>True</code> oder <code>False</code>) und viele weitere, insbesondere selbstdefinierte Typen.

```python
f=1.2345
b=True
type(f),type(b)
```

Natürlich gibt's die üblichen arithmetischen Operatoren: +-*/

```python
(1+2-3)*4/5
```

Fürs Potenzieren gibt's denn <code>**</code>-Operator: 

```python
2**10
```

Den Rest beim Teilen (mod) erhält man mit %, Ganzzahlteilung funktioniert mit //. 

```python
11%3,11//3
```

Was die Operatoren bewirken, hängt vom Datentyp ab. ZB können wir auch Strings addieren.

```python
first_string="python futtert"
second_string="Matlab zum Frühstück"
result=first_string+" "+second_string
print(result)
```

Zum Vergleichen gibt's ==,!=,>=,<=,<,>. Das Ergebnis ist entweder <code>True</code> oder <code>False</code>.

```python
x=12
y=13
x==y
```

```python
some_string="fortytwo"
some_string=='fortytwo'
```

```python
12!=12
```

```python
1234<=234
```

Die logischen Verknüpfungen <code>and</code> und <code>or</code> erlauben kompliziertere Ausdrücke.

```python
(x!=y) and (some_string=='fortytwo')
```

Kommentare werden mit # gemacht. Alles in einer Zeile nach # ist ein Kommentar und wird ignoriert.

```python
some_variable=12.23 # dies ist ein Kommentar 

# Eine ganze Zeile kann auch ein Kommentar sein 
### Oft setze ich gleich drei #. Das macht keinen Unterschied und ist rein kosmetisch.
```

### Aufgabe: Rechnung
Berechnen Sie, welchen Rest
$
2^{23}-7
$ geteilt durch 33 hinterlässt. Speichern Sie das Resultat in einer Variable namens <code>rest</code>.

<!-- grade_id: cell-b532a3e7ad38c504, points: 0, grade: false, solution: true -->
```python
### INSERT YOUR CODE HERE
raise NotImplementedError()
### END OF INSERT

rest
```

<!-- grade_id: cell-c3f289c482817820, points: 3, grade: true, solution: false -->
```python
### Dies ist eine TESTZELLE. 
### Wenn die Ausführung zu einem Fehler führt, haben Sie etwas falsch gemacht
### Wenn bei Ausführung nichts passiert, dann ist Ihre Lösung richtig oder zumindest nicht grob falsch
assert rest >=0
assert rest <=32

```